Protein-Induced Membrane Strain Drives Supercomplex Formation

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

This paper by Poverlein et al reports the substantial membrane deformation around the oxidative phosphorylation super complex, proposing that this deformation is a key part of super complex formation. I found the paper interesting and well-written.

We thank the Reviewer for finding our work interesting.

Analysis of the bilayer curvature is challenging on the fine lengthscales they have used and produces unexpectedly large energies (Table 1). Additionally, the authors use the mean curvature (Eq. S5) as input to the (uncited, but it seems clear that this is Helfrich) Helfrich Hamiltonian (Eq. S7). If an errant factor of one half has been included with curvature, this would quarter the curvature energy compared to the real energy, due to the squared curvature.

We thank the Reviewer for raising this important issue. We have now clarified in the SI and main manuscript that we employ the Helfrich model. In our initial implementation, we indeed used the mean curvature H, thereby missing a factor of 2. As the Reviewer correctly noted, this resulted in curvature deformation energies that were underestimated by a factor of ~4. We have now corrected for this effect in the revised analysis, and the updated Table 1. Importantly, however, this correction does not alter the general conclusions of our work that supercomplex formation relieves membrane strain and stabilizes the system. We have added an additional paragraph where we discuss the magnitude of the observed bending effects, and compared the previous estimates in literature:

SI:

“The local mean curvature of the membrane midplane was computed using the Helfrich model (4,5) …”

(4) W. Helfrich, Elastic properties of lipid bilayers theory and possible experiments. Zeitschrift für Naturforschung 28c, 693-703 (1973).

(5) F. Campelo et al., Helfrich model of membrane bending: From Gibbs theory of liquid interfaces to membranes as thick anisotropic elastic layers. Advances in Colloid and Interface Science 208, 25-33 (2014).

Main Text:

“which measures the energetic cost of deforming the membrane from a flat geometry (ΔGcurv) based on the Helfrich model (45, 46). …

Our analysis suggests that both contributions are substantially reduced upon formation of the SC, with the curvature penalty decreasing by 79.2 ± 5.2 kcal mol-1 (for a membrane area of ca. 1000 nm2) and the thickness penalty by 2.8 ± 2.0 kcal mol-1 (Table 1).”

“We note that the magnitude of the estimated bending energies (~10² kcal mol-1) (Table 1), while seemingly high at first glance, falls within the range expected for large-scale membrane deformation processes induced by large multi-domain proteins. For example, the Piezo mechanosensitive channel performs roughly 150kBT (≈ 90 kcal mol⁻¹) of work to bend the bilayer into its dome-like shape (65). Comparable energies have also been estimated for the nucleation of small membrane pores (66), while vesicle formation typically requires bending energies on the order of 300 kcal mol-1, largely independent of vesicle size (67). When normalized by the affected membrane area (~1000 nm2), these values correspond to an energy density of approximately 0.1 kcal mol-1 nm-2, which places our estimates within a biophysically reasonable regime. Notably, cryo-EM structures of several supercomplexes shows that such assemblies can impose significant curvature on the surrounding bilayer (36, 50, 68), supporting the notion that respiratory chain organization is closely coupled to local membrane deformation. Nevertheless, we expect that the absolute deformation energies may be overestimated, as the continuum Helfrich model neglects molecular-level effects such as lipid tilt and local rearrangements, which can partially relax curvature stresses and reduce the effective bending penalty near protein–membrane interfaces (69, 70).”

The bending modulus used (ca. 5 kcal/mol) is small on the scale of typically observed biological bending moduli. This suggests the curvature energies are indeed much higher even than the high values reported. Some of this may be due to the spontaneous curvature of the lipids and perhaps the effect of the protein modifying the nearby lipids properties.

The SI initially included an incorrect value for the bending modulus (20 kJ mol-1 instead of 20kBT), which has now been corrected. The revised value is consistent with experimentally reported bending moduli from X-ray scattering measurements, although there remains substantial uncertainty in the precise values across different experimental and computational studies.

“The bending deformation energy was computed from the mean curvature field H(x,y), assuming a constant bilayer bending modulus κ (taken as 20kbT = 11.85 kcal mol-1 (6)):”

(6) S. Brown et al., Comparative analysis of bending moduli in one-component membranes via coarsegrained molecular dynamics simulations. Biophysical Journal 124, 1–13 (2025).

It is unclear how CDL is supporting SC formation if its effect stabilizing the membrane deformation is strong or if it is acting as an electrostatic glue. While this is a weakenss for a definite quantification of the effect of CDL on SC formation, the study presents an interesting observation of CDL redistribution and could be an interesting topic for future work.

We agree with the Reviewer that future studies would be important to investigate the relationship between CDL-induced stabilization of membrane and its electrostatic effects.

In summary, the qualitative data presented are interesting (especially the combination of molecular modeling with simpler Monte Carlo modeling aiding broader interpretation of the results). The energies of the membrane deformations are quite large. This might reflect the roles of specific lipids stabilizing those deformations, or the inherent difficulty in characterizing nanometer-scale curvature.

We thank the Reviewer for appreciating our work and for the help in further improving our findings.

Reviewer #3 (Public review):

Summary:

In this contribution, the authors report atomistic, coarse-grained and lattice simulations to analyze the mechanism of supercomplex (SC) formation in mitochondria. The results highlight the importance of membrane deformation as one of the major driving forces for the SC formation, which is not entirely surprising given prior work on membrane protein assembly, but certainly of major mechanistic significance for the specific systems of interest.

We thank Reviewer 3 for appreciating the importance of our study.

Strengths:

The combination of complementary approaches, including an interesting (re)analysis of cryo-EM data, is particularly powerful, and might be applicable to the analysis of related systems. The calculations also revealed that SC formation has interesting impacts on the structural and dynamical (motional correlation) properties of the individual protein components, suggesting further functional relevance of SC formation. In the revision, the authors further clarified and quantified their analysis of membrane responses, leading to further insights into membrane contributions. They have also toned down the decomposition of membrane contributions into enthalpic and entropic contributions, which is difficult to do. Overall, the study is rather thorough, highly creative and the impact on the field is expected to be significant.

Weaknesses:

Upon revision, I believe the weakness identified in previous work has been largely alleviated.

We thank the Reviewer for their previous remarks, which allowed us to significantly improve our manuscript.

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

Mitochondrial membranes harbor the electron transport chain (ETC) that powers oxidative phosphorylation (OXPHOS) and drives the synthesis of ATP. Yet, under physiological conditions, the OXPHOS proteins operate as higher-order supercomplex (SC) assemblies, although their functional role remains poorly understood and much debated. By combining large-scale atomistic and coarse-grained molecular simulations with analysis of cryo-electron microscopic data and statistical as well as kinetic models, we show here that the formation of the mammalian I/III2 supercomplex reduces the molecular strain of inner mitochondrial membranes by altering the local membrane thickness and leading to an accumulation of both cardiolipin and quinone around specific regions of the SC. We find that the SC assembly also affects the global motion of the individual ETC proteins with possible functional consequences. On a general level, our findings suggest that molecular crowding and strain effects provide a thermodynamic driving force for the SC formation, with a possible flux enhancement in crowded biological membranes under constrained respiratory conditions.

Introduction

Biological electron transport chains (ETC) comprise a series of membrane-bound enzyme complexes (CI-CIV) that transfer electrons towards oxygen and protons across a biological membrane, creating a proton motive force (PMF) that powers the synthesis of ATP and active transport (1, 2). Biological membranes are often envisaged in the light of the classical fluid mosaic model (3), where the membrane proteins and lipids independently diffuse as a 2-dimensional solution. However, the biological membranes are highly crowded (4-6), particularly the inner mitochondrial membrane (IMM), with a protein content of around 45% (7). In this regard, the majority of the respiratory complexes do not diffuse independently within the membrane but operate rather as larger supercomplex (SC) assemblies, as first revealed by blue native gels (8) and confirmed by structural analyses (9-11). Moreover, several lipid molecules, particularly cardiolipin, a central anionic lipid of the IMM, is often tightly bound to the SCs (12-15). Yet, despite these structural insights, the functional role of the SC assemblies and the physical principles leading to their formation remain elusive and much debated.

SCs have been suggested to provide kinetic advantages (12, 16), favor substrate channeling ((10, 17), but cf. (18)), decrease the formation of reactive oxygen species (ROS) (19, 20), stabilize individual proteins (21-24), and/or reduce non-specific protein-protein interactions (18). It has also been recognized that the high protein density in membranes may play an important role in the SC formation (18), and influence, e.g., the lipid curvature (7, 25). On a physiological level, decreased SC formation has been linked to various patho-physiological conditions such as diabetes (26), heart failure (27), and apoptosis (28), although recent experiments (29) on mice unable to form SC showed no significant differences relative to WT mice under the studied conditions.

In recent years, cryo-electron microscopic (cryo-EM) and cryo-electron tomographic (cryo-TEM) studies revealed the structure of several SCs, including the mammalian 1.5 MDa SCI/III2 and the 1.7 MDa respirasome (SCI/III2/IV) (30-34), as well as various exotic megacomplexes isolated from single-celled eukaryotes (35, 36). SCs are also essential for some bacteria, such as the SCIII2/IV2 of actinobacteria (15, 37-39) that catalyze quinol oxidation coupled to oxygen reduction in a tight obligate protein assembly. Inhibition of such bacterial SCs provides potential avenues for treating pathogenic infections, such as tuberculosis, with emerging multi-drug resistant strains.

To probe the physical principles leading to the SC formation and its molecular consequences, we study here the structure and dynamics of the mammalian mitochondrial SCI/III2 by combining large-scale atomistic and coarse-grained molecular dynamics simulations with analysis of cryo-EM data and mathematical models that provide insight into crowding effects, membrane strain, deformation effects, as well as dynamic protein-lipid interactions influence SC formation.

The studied SCI/III2 catalyzes an NADH-driven quinone / cytochrome c reduction and offers a unique system to probe how the quinone (Q) substrate is transported within the SC assembly, and how the individual proteins interact within the membrane. In this regard, the quinol (QH2) produced by the CI module of the SC is re-oxidized by the CIII module, whilst the proton transport activity (CI: 2H+/e−; CIII2: 1H+/e−) powers the synthesis of ATP (Fig. 1). Interestingly, prior cryo-EM studies of this SC (40) revealed an asymmetric protein assembly with a trapped quinone in the proximal/distal Qo site, suggesting possible functional consequences. However, the dynamics and interaction of the SC and its substrates within the surrounding membrane environment and the individual active site remain poorly understood.

Structure and function of Complex I, Complex III2, and the SCI/III2.

(A) The structure of the OXPHOS complexes with lipid headgroups from atomistic MD simulations. (B, C) Overview of the structure and function of redox-driven proton pumping in CI and the Q-cycle of CIII2. (D-F) Molecular models of the (D) SCI/III2, (E) CI, and (F) CIII2 used in the atomistic MD simulations (water and ions were omitted for visual clarity).

Results

Protein-membrane interactions modulate the lipid dynamics and membrane strain

To probe how the SC formation affects the dynamics of the OXPHOS proteins, the lipid membrane, and the quinone pool, we studied the ovine SCI/III2 by both atomistic (aMD) and coarse-grained molecular dynamics (cgMD) simulations. While the aMD simulations give insight into the microsecond dynamics (11 μs in total, Table S1) of the system with atomic details, our cgMD simulations allowed us to probe the SC dynamics and its interactions in the membrane on much longer timescales (∼0.3 ms in cumulative simulation time, 50-75 μs/simulation) and with larger membrane size at a more approximate level (see Methods). The SC models were further compared against simulations of the individual CI and CIII2 complexes, as well as simulations of a lipid membrane, with a POPC:POPE:CDL ratio (2:2:1) as in the IMM. In the following, we define membrane strain as the local perturbations of the lipid bilayer induced by protein-membrane interactions. These include changes in (i) membrane thickness,(ii) the local membrane composition, (iii) lipid chain configurations, and (iv) local curvature of the membrane plane relative to an undisturbed, protein-free bilayer. Together, these phenomena reflect the thermodynamic effects associated with accommodating large protein complexes within the membrane.

Earlier version: We find that the individual OXPHOS complexes, CI and CIII2, induce pronounced membrane strain effects, supported both by our aMD (Fig. S2A) and cgMD simulations with a large surrounding membrane (Fig. 2G). These effects persist for both individual complexes and assembled SC. We observe a local decrease in the membrane thickness at the protein-lipid interface (Fig. 2G, Fig S2A,D,E), likely arising from the thinner hydrophobic belt region of the OXPHOS proteins (ca. 30 Å, Fig. S1A) relative to the lipid membrane (40.5 Å, Fig. S1D). We further observe ∼30% accumulation of cardiolipin at the thinner hydrophobic belt regions (Fig. 2H, Fig. S2B,F,G), with an inhomogeneous distribution around the OXPHOS complexes. While specific interactions between CDL and protein residues may contribute to this enrichment (Fig. 2N), CDL thermodynamically prefers thinner membranes (∼38 Å, Fig. S1D, Fig. S5E). These changes are further reflected in the reduced end-to-end distance of lipid chains in the local membrane belt (see Methods, Fig. S6, cf. also Refs. (41-44)). In addition to the perturbations in the local membrane thickness, the OXPHOS proteins also induce a subtle inward curvature towards the protein-lipid interface (Fig. S5G), which could modulate the accessibility of the Q/QH2 substrate into the active sites of CI and CIII2 (see below, section Discussion). This curvature is accompanied by a distortion of the local membrane plane itself (Fig. 2A-F, Fig. S4A-C, Fig. S8), with perpendicular leaflet displacements reaching up to ∼2 nm relative to the average leaflet plane.

Revised version: We find that the individual OXPHOS complexes, CI and CIII2, induce pronounced membrane strain effects, supported both by our aMD (Fig. S2A) and cgMD simulations with a large surrounding membrane (Fig. 2G). These effects persist for both individual complexes and assembled SC. We observe a local decrease in the membrane thickness at the protein-lipid interface (Fig. 2G, Fig S2A,D,E), likely arising from the thinner hydrophobic belt region of the OXPHOS proteins (ca. 30 Å, Fig. S1A) relative to the lipid membrane (40.5 Å, Fig. S1D). We further observe ∼30% accumulation of cardiolipin at the thinner hydrophobic belt regions (Fig. 2H, Fig. S2B,F,G), with an inhomogeneous distribution around the OXPHOS complexes. While specific interactions between CDL and protein residues may contribute to this enrichment (Fig. 2N), CDL also thermodynamically prefers thinner membranes (∼38 Å, Fig. S1D, Fig. S5E). These changes are further reflected in the reduced end-to-end distance of lipid chains in the local membrane belt (see Methods, Fig. S6, cf. also Refs. (41-44)). In addition to the perturbations in the local membrane thickness, the OXPHOS proteins also induce a subtle inward curvature towards the protein-lipid interface (Fig. S5G), which could modulate the accessibility of the Q/QH2 substrate into the active sites of CI and CIII2 (see below, section Discussion). This curvature is accompanied by a distortion of the local membrane plane itself (Fig. 2A-F, Fig. S4A-C, Fig. S8), with perpendicular leaflet displacements reaching up to ∼2 nm relative to the average leaflet plane.

SC formation affects the concentration of lipids and the membrane thickness.

(A-C) OXPHOS complexes viewed from the N-side of the membrane. (D-F) Membrane shift induced by (D) the SC, and (E,F) the isolated CI and CIII2. The membrane shift relative to the average membrane plane is colored by the shift in z-position viewed from the N-side. (G) Local membrane thickness around SC (top), CI and CIII2 (bottom). (H) Local cardiolipin concentration. (I) Local concentration of Q (Q+QH2). (J) Cryo-EM difference map. (K) Membrane shift determined from the cryo-EM map. (L) Membrane thickness determined from cryo-EM map. (M) Density of Q/QH2 (headgroup) from cgMD simulations. Insets: Protein-Q interactions from atomistic MD. (N) Density of cardiolipin (headgroup) from cgMD simulations. Insets: Protein-CDL interactions from atomistic MD.

Earlier version: To quantify the membrane strain effects, we analyzed the cgMD trajectories by projecting the membrane surface onto a 2-dimensional grid and calculating the local membrane height and thickness at each grid point. From these values, we quantified the local membrane curvature (Fig. S5H), which measures the energetic cost of deforming the membrane from a flat geometry (ΔGcurv). We also computed the energetics associated with changes in the membrane thickness, assessed from the deviations from an ideal local membrane in the absence of embedded proteins (ΔGthick, see Supporting Information, for technical details). Our analysis suggests that both contributions are substantially reduced upon formation of the SC, with the curvature penalty decreasing by 19.8 ± 1.3 kcal mol-1 and the thickness penalty by 2.8 ± 2.0 kcal mol-1 (Table 1). These results indicate a significant thermodynamic advantage for SC formation, as it minimizes lipid deformation and stabilizes the membrane environment surrounding Complex I and III.

Revised version: To quantify the membrane strain effects, we analyzed the cgMD trajectories by projecting the membrane surface onto a 2-dimensional grid and calculating the local membrane height and thickness at each grid point. From these values, we quantified the local membrane curvature (Fig. S5H), which measures the energetic cost of bending the membrane from a flat geometry (ΔGcurv) based on the Helfrich model (45, 46). We also computed the energetics associated with changes in the membrane thickness, assessed from the deviations from an ideal local membrane in the absence of embedded proteins (ΔGthick, see Supporting Information, for technical details). Our analysis suggests that both contributions are substantially reduced upon formation of the SC, with the curvature penalty decreasing by 79.2 ± 5.2 kcal mol-1 (for a membrane area of ca. 1000 nm2) and the thickness penalty by 2.8 ± 2.0 kcal mol-1 (Table 1). These results indicate a significant thermodynamic advantage for SC formation, as it minimizes lipid deformation and stabilizes the membrane environment surrounding Complex I and III.

Remarkably, the SC assembly partially releases this strain energy as a result of a smaller solvation area established at the SC-lipid interface (Fig. 2G, Fig. S5, Fig. S8, Table S4). We find that an increase in the lipid tail length decreases the relative stability of the SC (Fig. S5F), further supporting that the hydrophobic mismatch between the OXPHOS proteins and the lipid membrane modulates the strain effect (Fig. S5A-C).

Earlier version: Taken together, the analysis suggests that the OXPHOS complexes affect the mechanical properties of the membranes by inducing a small inwards curvature towards the protein-lipid interface (Fig. S5G,H), resulting in a membrane deformation effect, while the SC formation releases some deformation energy relative to the isolated OXPHOS complexes. The localization of specific lipids around the membrane proteins, as well as local membrane perturbation effects, is also supported by simulations of other membrane proteins (45, 46), suggesting that the effects could arise from general protein-membrane interactions.

Revised version: Taken together, the analysis suggests that the OXPHOS complexes affect the mechanical properties of the membranes by inducing a small inwards curvature towards the protein-lipid interface (Fig. S5G,H), resulting in a membrane deformation effect, while the SC formation releases some deformation energy relative to the isolated OXPHOS complexes. The localization of specific lipids around the membrane proteins, as well as local membrane perturbation effects, is also supported by simulations of other membrane proteins (47, 48), suggesting that the effects could arise from general protein-membrane interactions.

Earlier version: Although the current cryo-EM maps do not resolve the molecular details of the surrounding lipid membrane (but see below), we performed a spatial integration of the experimental cryo-EM density map by training a neural network model (see Methods) that allowed us to identify and calculate the membrane thickness of the lipid belt around Complex I (PDB ID: 6RFR, EMD-4873) (47). In this regard, we find that the experimental data also shows a statistically significant membrane thinning around CI relative to the SC at the protein-membrane interface (Fig. 2J-L), strongly supporting our simulation results. We note that during the finalization of this work, a membrane distortion effect was also reported for in situ high-resolution cryo-EM structures of respiratory SCs (48), thus providing additional support for our findings.

Revised version: Although the current cryo-EM maps do not resolve the molecular details of the surrounding lipid membrane (but see below), we performed a spatial integration of the experimental cryo-EM density map by training a neural network model (see Methods) that allowed us to identify and calculate the membrane thickness of the lipid belt around Complex I (PDB ID: 6RFR, EMD-4873) (49). In this regard, we find that the experimental data also shows a statistically significant membrane thinning around CI relative to the SC at the protein-membrane interface (Fig. 2J-L), strongly supporting our simulation results. We note that during the finalization of this work, a membrane distortion effect was also reported for in situ high-resolution cryo-EM structures of respiratory SCs (50), thus providing additional support for our findings.

We observe that the SC assembly influences the lipid and water dynamics at the protein-lipid interface and perturbs the local dielectric at the membrane plane (Fig. S1E). The dielectric constant (ε⊥) shows a local increase up to 1.5 nm from the membrane before it decays to the aqueous bulk dielectric (Fig. S1C,E). We note that this effect is strongly affected by the lipid type, particularly by the CDL that accumulates around the SC (Fig. 2H, Fig. S2B) and creates a micro-environment, which could enhance a local proton gradient.

Taken together, our combined findings suggest that the SC formation is affected by thermodynamic effects that reduce the molecular strain in the lipid membrane, whilst the perturbed micro-environment also affects the lipid and Q dynamics, as well as the dynamics of the OXPHOS proteins (see below).

The SC assembly alters the quinone dynamics but does not support substrate channeling

In addition to the phospholipids, the IMM contains 3-5% ubiquinone (Q), which functions as the electron carrier between CI and CIII2 in the form of ubiquinol (QH2). Our simulations, performed with ca. 5% Q/QH2 (1:1), suggest that the long isoprenoid tail of ubiquinone (Q10) has similar physico-chemical properties as the lipid tails, whilst the Q/QH2 headgroup is polar and localizes at the membrane surface, with a flip-flop rate between the leaflets of ca. 100-150 ns (Fig. S3A-C). The fast flip-flop motion could support the quinol diffusion from CI to CIII2, in which the quinol exits the negatively charged (N-side / matrix side) membrane surface of CI and enters the Qo site of CIII2 near the positively charged (P-side / intermembrane space) side of the membrane (see Discussion). Interestingly, the Q pool does not accumulate uniformly around the SC. Instead, we observe a ca. 3% increase in the local Q concentration near the substrate binding sites of both CI and CIII2 relative to the bulk membrane, whilst the immediate OXPHOS surroundings show an overall Q/QH2 depletion (Fig. 2I, Fig S2C, Fig S7D,E). This local Q pool could arise from specific interactions between Q and the OXPHOS proteins, e.g., in subunits ND1 and NDUFA9 of CI that contain several non-polar residues and surface charges that interact with the quinones (Fig. 2M, Fig. S14). In contrast, on the proximal side of CIII2, we observe a subtle decrease in the local Q concentration relative to CIII2 alone (Fig 2I, Fig S2C, Fig S7D,E), an effect that could arise from a shift in the membrane plane near the Qo site and affect the substrate access to the proximal Qo site (Fig S4, S8). These observations are consistent with the occupied distal Qo site and empty proximal Qo site observed in the cryo-EM structure of the SCI/III2 (40).

Overall, while our simulations indicate that quinones accumulate around the OXPHOS proteins, our data do not support substrate channeling (10) (see Discussion), as we observe neither a contiguous region of an elevated Q pool between the complexes, nor a directed diffusion pathway of the quinol from CI to CIII2 (Fig. 2I, Fig. S3D-I). Moreover, the locally hindered access to the proximal Qo site due to the membrane shift could further hamper the substrate binding into the active site (Fig. S4D-E).

The SC assembly alters the conformational dynamics of CI and CIII2

Earlier version: We next probed how the SC assembly affects the dynamics of the individual OXPHOS proteins based on our atomistic MD simulations of CI, CIII2, and the SCI/III2. Despite the large (1.5 MDa) SC, we note the SC assembly is stabilized by only a few specific hydrogen-bonding and charged interactions at the interface of NDUFB9/NDUFB4 and UQCRC1, particularly the Arg29FB4-Asp260C1/Glu259C1 ion-paired network that interacts by ca. -10 kcal mol-1 during the MD simulations (Fig. 3A,C). Indeed, deletion of this Glu258-Asp260 region in UQCRC1 led to a drastic reduction in the stability of SC in mice (29) (see Fig. S11), thus supporting the importance of these interactions in the SC.

Revised version: We next probed how the SC assembly affects the dynamics of the individual OXPHOS proteins based on our atomistic MD simulations of CI, CIII2, and the SCI/III2. Despite the large (1.5 MDa) SC, we note the SC assembly is stabilized by only a few specific hydrogen-bonding and charged interactions at the interface of NDUFB9/NDUFB4 and UQCRC1, particularly the Arg29FB4-Asp260C1/Glu259C1 ion-paired network that interacts by ca. -10 kcal mol-1 during the MD simulations (Fig. 3A,C). Indeed, deletion of this Glu258-Asp260 region in UQCRC1 led to a drasticreduction in the stability of SC in mice (29) (see Fig. S11), thus supporting the importance of these interactions in the SC.

Conformational dynamics of the OXHPOS proteins and the SC.

(A) Subunits comprising the SC interface. Inset: Interaction near the DED-loop. (B) Normal modes derived from essential dynamics analysis for the SC (left), CI (middle), and CIII2 (right). See SI Appendix, Movies 1-5 for normal modes of the SC, CI, and CIII2. (C) Decomposition of interaction energies within the SC. (D) Distribution of the opening/closing mode (left) and the twisting mode (right) for the isolated CI (in orange) and the CI within the SC (in blue). (E) Differences in the dynamics of CIII2 for minimum and maximum values of mode 1. (F) Distribution of the CI domain angle in the SC (in blue) and CI (in orange). (G,J) Conformational changes in the Rieske subunit between SC (in blue) and CIII2 (in green) affect (G) the QH2 binding in the proximal Qo site and (J) the Rieske FeS-heme c1 distance in the distal monomer. (H) Distribution of the dihedral angle in TM3 of ND6. (I) Overview of the SC showing the locations of ND6 and Rieske subunits, as well as the proximal and distal protomers of CIII2.

Earlier version: Based on essential dynamics analysis (EDA, see Methods, Movies S1-S5) that projects out dominant protein motion, we find that the CI and CIII2 modules of the SC undergo a back-and-forth rocking motion around the interface region, leading to a breathing motion between the hydrophilic domains of the OXPHOS complexes on the N-side of the membrane (matrix side) (Fig. 3B, Mode 1, Movie S1), while the second dominant motion (Mode 2, Movie S1) couples the opening/closing motion of CI (cf. also ref (49)) with a rocking motion of CIII2. The third mode arises from a combination of the twisting motion of the hydrophilic domain of CI and a minor rocking motion of CIII2 (Mode 3, Movie S1), which could influence the membrane accessibility of the Q sites. In this regard, our graph theoretical analysis (Fig. S11C,D) further indicates that ligand binding to Complex I induces a dynamic crosstalk between NDUFA5 and NDUFA10, consistent with previous work (50, 51), and affecting also the motion of UQCRC2 with respect to its surroundings. Taken together, these effects suggest that the dynamics of CI and CIII2 show some correlation that could result in allosteric effects, as also indicated based on cryo-EM analysis (40).

Revised version: Based on essential dynamics analysis (EDA, see Methods, Movies S1-S5) that projects out dominant protein motion, we find that the CI and CIII2 modules of the SC undergo a back-and-forth rocking motion around the interface region, leading to a breathing motion between the hydrophilic domains of the OXPHOS complexes on the N-side of the membrane (matrix side) (Fig. 3B, Mode 1, Movie S1), while the second dominant motion (Mode 2, Movie S1) couples the opening/closing motion of CI (cf. also ref (51)) with a rocking motion of CIII2. The third mode arises from a combination of the twisting motion of the hydrophilic domain of CI and a minor rocking motion of CIII2 (Mode 3, Movie S1), which could influence the membrane accessibility of the Q sites. In this regard, our graph theoretical analysis (Fig. S11C,D) further indicates that ligand binding to Complex I induces a dynamic crosstalk between NDUFA5 and NDUFA10, consistent with previous work (52, 53), and affecting also the motion of UQCRC2 with respect to its surroundings. Taken together, these effects suggest that the dynamics of CI and CIII2 show some correlation that could result in allosteric effects, as also indicated based on cryo-EM analysis (40).

Earlier version: We note that the SC assembly induces a subtle conformational change in CI that increases the angle between the hydrophilic and membrane domains (Fig. 3F). This motion is linked to the active / deactive-transition (A/D) (52), which regulates CI activity, particularly under hypoxic / anoxic conditions, and hinders reverse electron transfer (RET) (52-57) (but cf. also ref. (58)). These changes in CI lead to a subtle shift in the conformation of key transmembrane helices (TM) and surrounding conserved loops. In the SCI/III2, the TM3 of subunit ND6 undergoes a twist towards the deactive state (Fig. 3H, Fig. S13A-D), while the β1-β2 loop of subunit NDUFS2 exhibits lower flexibility in both the apo and the QH2 bound state relative to CI alone (Fig. S13E-H). These regions are central in modulating the coupling between the proton transport and the electron transfer activities ((59), cf. also Ref. (60)), and the dynamics of these regions could thus have functional consequences. CI also shows a dominant vibrational motion that twists the hydrophilic domain relative to the membrane domain (Fig. 3B,D), and samples a wider angle between hydrophilic and membrane domains in the SC, (Fig. 3F) that result in a larger twisting angle relative to the isolated CI (Fig. 3F). These alternations suggest that the opening/closing and twisting motions could be sterically hindered by the adjacent CIII2 (see Movies S1-S5), and that the active/deactive (A/D) transition is coupled to the global motions of the SC. Consistent with our previous findings (49), CI shows also another dominant bending motion, where the membrane domain rotates relative to the hydrophilic domain within the membrane plane.

Revised version: We note that the SC assembly induces a subtle conformational change in CI that increases the angle between the hydrophilic and membrane domains (Fig. 3F). This motion is linked to the active / deactive-transition (A/D) (54), which regulates CI activity, particularly under hypoxic / anoxic conditions, and hinders reverse electron transfer (RET) (54-59) (but cf. also ref. (60)). These changes in CI lead to a subtle shift in the conformation of key transmembrane helices (TM) and surrounding conserved loops. In the SCI/III2, the TM3 of subunit ND6 undergoes a twist towards the deactive state (Fig. 3H, Fig. S13A-D), while the β1-β2 loop of subunit NDUFS2 exhibits lower flexibility in both the apo and the QH2 bound state relative to CI alone (Fig. S13E-H). These regions are central in modulating the coupling between the proton transport and the electron transfer activities ((61), cf. also Ref. (62)), and the dynamics of these regions could thus have functional consequences. CI also shows a dominant vibrational motion that twists the hydrophilic domain relative to the membrane domain (Fig. 3B,D), and samples a wider angle between hydrophilic and membrane domains in the SC, that result in a larger twisting angle relative to the isolated CI (Fig. 3F). These alternations suggest that the opening/closing and twisting motions could be sterically hindered by the adjacent CIII2 (see Movies S1-S5), and that the active/deactive (A/D) transition is coupled to the global motions of the SC. Consistent with our previous findings (51), CI shows also another dominant bending motion, where the membrane domain rotates relative to the hydrophilic domain within the membrane plane.

Earlier version: CIII2 also undergoes dynamical changes upon the SC formation, particularly by enhancing the normal modes connected to the opening/closing transitions around the iron-sulfur (FeS) Rieske center (Fig. 3B,E) that could affect the electron transfer activity of CIII2 (61). In the SCI/III2, QH2 binds further away from the FeS center in the proximal Qo site, while in the distal site, the FeS center moves closer to the heme c1 (Fig 3G,I,J). This asymmetry suggests that the global motion of the SC could regulate the ‘preferred’ Qo site for the electron bifurcation process, and possibly favor the electron transfer onwards to the Complex IV (CIV) that resides on the distal side of the CIII2 module in the respirasome (SCI/III2/IV) (30-32) (cf. also graph theoretical analysis, Fig. S11C,D). Taken together, these effects suggest that the dynamics of CI and CIII2 show some correlation that could result in allosteric effects, as also suggested by the recent in situ cryo-EM study of mitochondrial SCs (40). In this regard, we find that the ligand state of CI (apo or QH2) affects the conformational dynamics and the interaction interface of the SCI/CIII2 (Fig. S12A-E). This surprising long-range effect is likely to result from the increased flexibility of the apo state (Fig. S12G) that, in turn, modulates the interaction at the interface of the SC. Interestingly, similar ligand dependent conformational changes affecting both the CI and CIII2 domains of the SC are also supported by recent in situ cryo-EM structures of mitochondrial SCs (48).

Revised version: CIII2 also undergoes dynamical changes upon the SC formation, particularly by enhancing the normal modes connected to the opening/closing transitions around the iron-sulfur (FeS) Rieske center (Fig. 3B,E) that could affect the electron transfer activity of CIII2 (63). In the SCI/III2, QH2 binds further away from the FeS center in the proximal Qo site, while in the distal site, the FeS center moves closer to the heme c1 (Fig 3G,I,J). This asymmetry suggests that the global motion of the SC could regulate the ‘preferred’ Qo site for the electron bifurcation process, and possibly favor the electron transfer onwards to the Complex IV (CIV) that resides on the distal side of the CIII2 module in the respirasome (SCI/III2/IV) (30-32) (cf. also graph theoretical analysis, Fig. S11C,D). Taken together, these effects suggest that the dynamics of CI and CIII2 show some correlation that could result in allosteric effects, as also suggested by the recent in situ cryo-EM study of mitochondrial SCs (40). In this regard, we find that the ligand state of CI (apo or QH2) affects the conformational dynamics and the interaction interface of the SCI/CIII2 (Fig. S12A-E). This surprising long-range effect is likely to result from the increased flexibility of the apo state (Fig. S12G) that, in turn, modulates the interaction at the interface of the SC. Interestingly, similar ligand dependent conformational changes affecting both the CI and CIII2 domains of the SC are also supported by recent in situ cryo-EM structures of mitochondrial SCs (50).

Enthalpy-entropy compensation drives SC formation

Earlier version: Our molecular simulations suggest that while the membrane strain provides a thermodynamic driving force for the SC formation, the molecular interactions at the interface of the assembly are essential for enthalpically stabilizing the SC over non-specific protein assemblies (Fig. S12). To assess how the SC formation is affected by the strain effects, protein-protein interactions, and the protein/lipid ratio, we developed a simple statistical-mechanical lattice model of the CI and CIII2 diffusion in the IMM (Fig. 4A). To this end, the protein-protein interactions were described by tunable interaction energies (Especific and Enon-specific), modeling the specific hydrogen-bond / salt-bridges interactions and the non-specific interaction, while the membrane-protein interactions were tuned by the strain energy (Estrain). Our model suggests that the ratio between the strain energy and specific interactions indeed provides the key driving force for the SC formation, with a small (30%) decrease in Estrain leading to a ca. 60% decrease in the SC population (Fig. 4B). Our model further predicts that the SC population drastically increases at a specific interaction energy threshold (Especific < -2 kBT) relative to the strain contribution (Fig. 4B), whilst the number of non-specific assemblies (adjacent proteins in non-SC orientations) relative to SCs is determined by the ratio of Especific to Enon-specific. With a high membrane strain, we find that an increase in temperature also favors the formation of SC assemblies, as it is entropically favored to reduce the overall number of ordered (“strained”) lipid molecules around the OXPHOS proteins (Fig. S15). Similarly, a high protein-to-lipid ratio significantly increases the population of SCs, suggesting that a high protein packing in IMMs favors the SC formation, whilst the strength of specific interactions determines the relative population of non-specific assemblies in crowded environments. Taken together, our lattice model, despite its simplicity, captures and validates key features observed in our molecular simulations and supports that the SC formation is affected by both enthalpic and entropic effects.

Revised version: Our molecular simulations suggest that while the membrane strain provides a thermodynamic driving force for the SC formation, the molecular interactions at the interface of the assembly are essential for enthalpically stabilizing the SC over non-specific protein assemblies (Fig. S12). To assess how the SC formation is affected by the strain effects, protein-protein interactions, and the protein/lipid ratio, we developed a simple statistical-mechanical lattice model of the CI and CIII2 diffusion in the IMM (Fig. 4A). To this end, the protein-protein interactions were described by tunable interaction energies (Especific and Enon-specific), modeling the specific hydrogen-bond / salt-bridge interactions and the non-specific interaction, while the membrane-protein interactions were tuned by the strain energy (Estrain). Our model suggests that the ratio between the strain energy and specific interactions indeed provides the key driving force for the SC formation, with a small (30%) decrease in Estrain leading to a ca. 60% decrease in the SC population (Fig. 4B). Our model further predicts that the SC population drastically increases at a specific interaction energy threshold (Especific < -2 kBT) relative to the strain contribution (Fig. 4B), whilst the number of non-specific assemblies (adjacent proteins in non-SC orientations) relative to SCs is determined by the ratio of Especific to Enon-specific. With a high membrane strain, we find that an increase in temperature also favors the formation of SC assemblies, as it is entropically favored to reduce the overall number of ordered (“strained”) lipid molecules around the OXPHOS proteins (Fig. S15). Similarly, a high protein-to-lipid ratio significantly increases the population of SCs, suggesting that a high protein packing in IMMs favors the SC formation, whilst the strength of specific interactions determines the relative population of non-specific assemblies in crowded environments. Taken together, our lattice model, despite its simplicity, captures and validates key features observed in our molecular simulations and supports that the SC formation is affected by both enthalpic and entropic effects.

Lattice model of SC formation and crowding effects in the IMM.

Earlier version: (A) Possible specific and non-specific interactions between CI (blue/white) and CIII2 (red) and their respective energies in the lattice model. (B) SC population as a function of the specific interaction energy (Especific) and molecular strain (Estrain). (C) Specific and non-specific assemblies as a function of the protein-lipid ratio with varying specific interaction energies (Especific). (D) Representative protein arrangement in crowded IMMs. (E) Average edge-to-edge distance distributions and nearest neighbor distance (red line) for specific protein-Q/protein contacts. (F) Nearest neighbor distance between Q and CI as a function of the Q/QH2 ratio. The dashed line indicates the 7 nm distance between active sites in the SC.

Revised version: (A) Possible specific and non-specific interactions between CI (blue/white) and CIII2 (red) and their respective energies in the lattice model. (B) SC population as a function of the specific interaction energy (Especific) and molecular strain (Estrain). (C) Specific and non-specific assemblies as a function of the protein-lipid ratio with varying specific interaction energies (Especific). (D) Representative protein arrangement in crowded IMMs. Average edge-to-edge distance distributions and nearest neighbor distance (red line) for specific protein-Q/protein contacts. Nearest neighbor distance between Q and CI as a function of the Q/QH2 ratio. The dashed line indicates the 7 nm distance between active sites in the SC.

Discussion

Earlier version: We have shown here that the respiratory chain complexes perturb the IMM by affecting the local membrane dynamics. The perturbed thickness and alteration in the lipid dynamics lead to an energetic penalty, which can be related to molecular strain effects, as suggested by the changes of both the internal energy of lipid and their interaction with the surroundings (Fig. S2, S5, S6), which are likely to be of enthalpic origin. However, lipid binding to the OXPHOS complex also results in a reduction in the translational and rotational motion of the lipids and quinone (Fig. S9, S10), which could result in entropic changes. The strain effects are therefore likely to arise from a combination of enthalpic and entropic effects. In this regard, we suggest that the SC assemblies form by condensation of local high-entropic membrane regions around the OXPHOS proteins (Fig. 5A). However, as the entropic effects also favor the formation of non-specific protein assemblies, unique interactions, such as the ion-paired network around UQCRC1 and NDUFB9/NDUFB4 (Fig. 3A), are likely to enthalpically stabilize the SC assemblies. The suggested principles show similarities to the hydrophobic effect driving protein folding by condensation of locally ordered water clusters around unfolded protein patches (62).

Revised version: We have shown here that the respiratory chain complexes perturb the IMM by affecting the local membrane dynamics. The perturbed thickness and alteration in the lipid dynamics lead to an energetic penalty, which can be related to molecular strain effects, as suggested by the changes of both the internal energy of lipid and their interaction with the surroundings (Fig. S2, S5, S6), which are likely to be of enthalpic origin. However, lipid binding to the OXPHOS complex also results in a reduction in the translational and rotational motion of the lipids and quinone (Fig. S9, S10), which could result in entropic changes. The strain effects are therefore likely to arise from a combination of enthalpic and entropic effects. In this regard, we suggest that the SC assemblies form by condensation of local high-entropic membrane regions around the OXPHOS proteins (Fig. 5A). However, as the entropic effects also favor the formation of non-specific protein assemblies, unique interactions, such as the ion-paired network around UQCRC1 and NDUFB9/NDUFB4 (Fig. 3A), are likely to enthalpically stabilize the SC assemblies. The suggested principles show similarities to the hydrophobic effect driving protein folding by condensation of locally ordered water clusters around unfolded protein patches (64).

Proposed thermodynamic and functional effects of SCs.

(A) The individual OXPHOS complexes, CI and CIII2 induce prominent molecular strain in the surrounding lipid membrane (dark green area) due to a hydrophobic mismatch between the protein and the membrane. Condensation of locally strained membrane patches entropically drives the SC formation, and leads to an overall reduction of the membrane strain (light green area around the SC), favoring the accumulation of cardiolipin around the SC. (B) During normal respiratory conditions, each OXPHOS complex is surrounded by multiple (ca. 6) quinol/quinone molecules that can act as substrates for the proteins. At limiting QH2 concentration, the quinol diffusion between CI and CIII2 becomes rate-limiting, and leads to a kinetic advantage of the SC (see also Fig. 4F).

Added in the revised version: We note that the magnitude of the estimated bending energies (∼102 kcal mol−1) (Table 1), while seemingly high at first glance, falls within the range expected for large-scale membrane deformation processes induced by large multi-domain proteins. For example, the Piezo mechanosensitive channel performs roughly 150kBT (≈ 90 kcal mol−1) of work to bend the bilayer into its dome-like shape (65). Comparable energies have also been estimated for the nucleation of small membrane pores (66), while vesicle formation typically requires bending energies on the order of 300 kcal mol−1, largely independent of vesicle size (67). When normalized by the affected membrane area (∼1000 nm2), these values correspond to an energy density of approximately 0.1 kcal mol−1 nm−2, which places our estimates within a biophysically reasonable regime. Notably, cryo-EM structures of supercomplexes shows that such assemblies can impose significant curvature on the surrounding bilayer (36, 50, 68), supporting the notion that respiratory chain organization is closely coupled to local membrane deformation. Nevertheless, we expect that the absolute deformation energies may be overestimated, as the continuum Helfrich model neglects molecular-level effects such as lipid tilt and local rearrangements, which can partially relax curvature stresses and reduce the effective bending penalty near protein–membrane interfaces (69, 70).

Earlier version: We further probed the thermodynamic effects underlying the SC formation by developing a 2D lattice model. Despite its simplicity, the model supports that the SC stability is determined by a delicate balance between membrane strain effects and specific protein-protein interactions, but also strongly affected by the protein concentration and temperature effects (Fig. 4B). Moreno-Loshuertos et al. (63) recently suggested that elevated temperatures (>43°C) may indeed disrupt SCs, although the stability of the individual OXPHOS proteins was also decreased in the studied conditions. As the enthalpic effects are of electrostatic origin, the SC stability could also be sensitive to the ionic strength and the PMF, which in turn depends on the metabolic state of the mitochondria.

Revised version: We further probed the thermodynamic effects underlying the SC formation by developing a 2D lattice model. Despite its simplicity, the model supports that the SC stability is determined by a delicate balance between membrane strain effects and specific protein-protein interactions, but also strongly affected by the protein concentration and temperature effects (6Fig. 4B). Moreno-Loshuertos et al. (71) recently suggested that elevated temperatures (>43°C) may indeed disrupt SCs, although the stability of the individual OXPHOS proteins was also decreased in the studied conditions. As the enthalpic effects are of electrostatic origin, the SC stability could also be sensitive to the ionic strength and the PMF, which in turn depends on the metabolic state of the mitochondria.

Earlier version: At the molecular level, the SC formation leads to an accumulation of CDL at the protein-membrane interface (Fig. 2N), as well as a local Q / QH2 pool near the substrate channel of CI and CIII2 (Fig. 2M). We find that CDL prefers thinner membranes relative to the neutral phospholipids (PE/PC, Fig. S5E), and could thus partially compensate for the hydrophobic mismatch between the OXPHOS proteins and the membrane (Fig. S1). The Q diffusion is affected by both specific interactions with OXPHOS proteins, as well as the local membrane thickness (Fig. 2I). The CDL around the SC is thus likely to have both structural and functional consequences, consistent with its effect on the activity and dynamics of several membrane proteins ((64-67), cf. also Ref. (49)). In this regard, CDL was suggested to enhance the substrate dynamics within the Q-tunnel of CI that requires a twisting-bending motion around the membrane and hydrophilic domains (49), whilst destabilization of CDL-binding sites has indeed been shown to disrupt SCs, e.g., the SCIII2IV1 in yeast (12). Moreover, defects in the CDL synthesis, e.g., in the Barth syndrome, (68), result in the disassembly of SC, indirectly supporting the involvement of CDL as a “SC glue”. In this regard, electrostatic effects arising from the negatively charged cardiolipin headgroup could play an important role in the interaction of the OXPHOS complexes. While CDL was modeled here in the singly anionic charged state (but cf. Fig. S5E), we note that the local electrostatic environment could tune their pKa that result in protonation changes of the lipid, consistent with its function as a proton collecting antenna (69).

Revised version: At the molecular level, the SC formation leads to an accumulation of CDL at the protein-membrane interface (Fig. 2N), as well as a local Q / QH2 pool near the substrate channel of CI and CIII2 (Fig. 2M). We find that CDL prefers thinner membranes relative to the neutral phospholipids (PE/PC, Fig. S5E), and could thus partially compensate for the hydrophobic mismatch between the OXPHOS proteins and the membrane (Fig. S1). The Q diffusion is affected by both specific interactions with OXPHOS proteins, as well as the local membrane thickness (Fig. 2I). The CDL around the SC is thus likely to have both structural and functional consequences, consistent with its effect on the activity and dynamics of several membrane proteins ((72-75), cf. also Ref. (51)). In this regard, CDL was suggested to enhance the substrate dynamics within the Q-tunnel of CI that requires a twisting-bending motion around the membrane and hydrophilic domains (51), whilst destabilization of CDL-binding sites has indeed been shown to disrupt SCs, e.g., the SCIII2IV1 in yeast (12). Moreover, defects in the CDL synthesis, e.g., in the Barth syndrome, (76), result in the disassembly of SC, indirectly supporting the involvement of CDL as a “SC glue”. In this regard, electrostatic effects arising from the negatively charged cardiolipin headgroup could play an important role in the interaction of the OXPHOS complexes. While CDL was modeled here in the singly anionic charged state (but cf. Fig. S5E), we note that the local electrostatic environment could tune their pKa that result in protonation changes of the lipid, consistent with its function as a proton collecting antenna (77).

Earlier version: In addition to the changes in the membrane properties, we observed that the SC formation modulates the conformational dynamics of the individual OXPHOS proteins, especially the large-scale bending-twisting motion of CI, the conformation of individual TM helices and conserved loops around proton channels in CI, as well as the motion of the Rieske domain in CIII2 (Fig. 3, Fig. S12). The dynamics of these regions are likely to module the activity of the OXPHOS proteins, e.g., the A/D transition of CI (57, 70) that regulate the ΔpH-driven quinol oxidation and reverse electron transfer (71, 72). Indeed, blocking loop motions surrounding these regions inhibits the proton pumping activity of CI (73), whilst the perturbed motion of the Rieske domain could modulate the electron bifurcation in CIII2 and subsequent electron transfer to CIV. In this regard, we suggest the preferred QH2 binding in the distal Qo site has functional implications for the respirasome (SCI/III2/IV), where the CIV module is located on the distal side of the CIII2 protomer. The changes in the conformational dynamics upon SC formation may thus affect ROS production (17, 74) via the A/D transition of CI (57), although it should be noted that no differences in ROS generation were observed for mice unable to form SCs relative to WT mice (with ca. 75% SCs) under normal conditions (29). It is possible that differences occur only under more strained conditions, e.g., in hypoxia or together with disease-related mutations in the OXPHOS proteins, where the SCs could become functionally more important (see below).

Revised version: In addition to the changes in the membrane properties, we observed that the SC formation modulates the conformational dynamics of the individual OXPHOS proteins, especially the large-scale bending-twisting motion of CI, the conformation of individual TM helices and conserved loops around proton channels in CI, as well as the motion of the Rieske domain in CIII2 (Fig. 3, Fig. S12). The dynamics of these regions are likely to module the activity of the OXPHOS proteins, e.g., the A/D transition of CI (59, 78) that regulate the ΔpH-driven quinol oxidation and reverse electron transfer (79, 80). Indeed, blocking loop motions surrounding these regions inhibits the proton pumping activity of CI (81), whilst the perturbed motion of the Rieske domain could modulate the electron bifurcation in CIII2 and subsequent electron transfer to CIV. In this regard, we suggest the preferred QH2 binding in the distal Qo site has functional implications for the respirasome (SCI/III2/IV), where the CIV module is located on the distal side of the CIII2 protomer. The changes in the conformational dynamics upon SC formation may thus affect ROS production (17, 82) via the A/D transition of CI (59), although it should be noted that no differences in ROS generation were observed for mice unable to form SCs relative to WT mice (with ca. 75% SCs) under normal conditions (29). It is possible that differences occur only under more strained conditions, e.g., in hypoxia or together with disease-related mutations in the OXPHOS proteins, where the SCs could become functionally more important (see below).

Earlier version: To understand how SCs could influence the charge currents in the IMMs, we note that the local increase of the quinol concentration near the active site of CI, and the local decreased quinol pool around the proximal Qo site of CIII2 could create a substrate gradient (∇c) and affect the quinol flux between CI and CIII2, JCI→CIII = -Dmem∇c (Dmem – Q/QH2 diffusion constant) – if the quinol concentration is rate-limiting for the function of CI or CIII2 (but see below). The 2D-diffusion time for the quinol, τ = <r2>/4Dmem, between the OXPHOS complexes depends on the effective protein-protein distance (r), which can be estimated from the protein packing density (cf. (7, 34, 75-79)). Using the experimental protein copy numbers in IMM (cf. Refs. (7, 34, 75-79) and Extended SI Text, Table S5), we obtain average edge-to-edge distance between CI and CIII2 of around 12 nm (Fig. 4E, cf. also (7)), which can be compared to the edge-to-edge distance of ca. 7 nm / 12 nm between the Q tunnel of CI and the proximal / distal Qo sites of CIII2 within the SC. The reduced diffusion distance could thus provide a subtle rate enhancement (τSC(r=7 nm)/τnonSC(r=12 nm)∼0.3) in favor of the SCs under substrate-limited conditions (Fig. 4E-F). We expect that the small variations in the lateral quinol diffusion (0.4 nm ns-1 near CI, 0.8 nm ns-1 in bulk) around the SC (Fig S3D-I, Fig S9) could favor the diffusion along the membrane patches, although this diffusion is also affected by non-specific collisions. However, due to the overall 3-5% Q / QH2 concentration in IMMs, we estimate that each OXPHOS protein is surrounded by around 6 quinone / quinol species depending on the respiratory conditions (Fig. 5, Extended SI Text), and leading to a small (0.3 nm) nearest neighbor distance between Q and CI/CIII2 (Fig. 4E, see SI Appendix). This implies that each OXPHOS protein has a saturated substrate Q pool within its “reaction sphere” upon high charge flux conditions (here 50% reduced Q pool). In other words, the CIII2 node of the SC does not need to rely on the quinol generated at the CI node, but can instead oxidize a quinol molecule from the surrounding Q pool. In this regard, Hirst and co-workers (80) found that the alternative oxidase (AOX) co-reconstituted into liposomes outcompetes the quinol re-oxidation rate of the SCI/III2, suggesting that the diffusion of the Q-species between the OXPHOS complexes is not rate-limiting under the studied conditions. However, at low QH2/Q ratios (<5% in the model), the minimal nearest neighbor distance between the Q pool and CI/CIII2 drastically increases (Fig. 4D-F) that in turn leads to a kinetic preference for the SC (Fig. 4D).

Revised version: To understand how SCs could influence the charge currents in the IMMs, we note that the local increase of the quinol concentration near the active site of CI, and the local decreased quinol pool around the proximal Qo site of CIII2 could create a substrate gradient (∇c) and affect the quinol flux between CI and CIII2, JCI→CIII = -Dmem∇c (Dmem – Q/QH2 diffusion constant) – if the quinol concentration is rate-limiting for the function of CI or CIII2 (but see below). The 2D-diffusion time for the quinol, τ = <r2>/4Dmem, between the OXPHOS complexes depends on the effective protein-protein distance (r), which can be estimated from the protein packing density (cf. (7, 34, 83-87)). Using the experimental protein copy numbers in IMM (cf. Refs. (7, 34, 83-87) and Extended SI Text, Table S5), we obtain average edge-to-edge distance between CI and CIII2 of around 12 nm (Fig. 4E, cf. also (7)), which can be compared to the edge-to-edge distance of ca. 7 nm / 12 nm between the Q tunnel of CI and the proximal / distal Qo sites of CIII2 within the SC. The reduced diffusion distance could thus provide a subtle rate enhancement (τSC(r=7 nm)/τnonSC(r=12 nm)∼0.3) in favor of the SCs under substrate-limited conditions (Fig. 4E-F). We expect that the small variations in the lateral quinol diffusion (0.4 nm ns-1 near CI, 0.8 nm ns-1 in bulk) around the SC (Fig S3D-I, Fig S9) could favor the diffusion along the membrane patches, although this diffusion is also affected by non-specific collisions. However, due to the overall 3-5% Q / QH2 concentration in IMMs, we estimate that each OXPHOS protein is surrounded by around 6 quinone / quinol species depending on the respiratory conditions (Fig. 5, Extended SI Text), and leading to a small (0.3 nm) nearest neighbor distance between Q and CI/CIII2 (Fig. 4E, see SI Appendix). This implies that each OXPHOS protein has a saturated substrate Q pool within its “reaction sphere” upon high charge flux conditions (here 50% reduced Q pool). In other words, the CIII2 node of the SC does not need to rely on the quinol generated at the CI node, but can instead oxidize a quinol molecule from the surrounding Q pool. In this regard, Hirst and co-workers (88) found that the alternative oxidase (AOX) co-reconstituted into liposomes outcompetes the quinol re-oxidation rate of the SCI/III2, suggesting that the diffusion of the Q-species between the OXPHOS complexes is not rate-limiting under the studied conditions. However, at low QH2/Q ratios (<5% in the model), the minimal nearest neighbor distance between the Q pool and CI/CIII2 drastically increases (Fig. 4D-F) that in turn leads to a kinetic preference for the SC (Fig. 4D).

At the microscopic level, the SC influences the lipid and water dynamics at the protein-lipid interface and the physico-chemical properties of the IMM, including its local dielectric properties of the membrane. It is possible that such local micro-environments have functional implications that could affect the proton conduction along the membrane, with important bioenergetic consequences. Taken together, our combined findings suggest that SC forms as a result of a complex interplay between molecular interactions and membrane strain effects that control the functional dynamics of the OXPHOS proteins.

Materials and Methods

Atomistic MD Simulations

Earlier version: Atomistic molecular dynamics (aMD) simulations of the ovine Complex I (CI), Complex III2 (CIII2), and the I-III2 supercomplex (SCI/III2) were performed in a POPC:POPE:cardiolipin (2:2:1) membrane containing 5 mol% QH2 / Q (1:1 ratio). Cardiolipin was modeled as tetraoleoyl cardiolipin (18:1/18:1/18:1/18:1) with a headgroup modeled in a singly protonated state (with Qtot=-1). The system was solvated with TIP3P water molecules and 150 mM NaCl. The simulations were performed in different ligand states (apo, Q/QH2 bound, Table S1) at T=310 K and p=1 bar using NAMD2.14 (81) with a 2 fs integration timestep and long-range electrostatic interactions treated using the Particle Mesh Ewald approach. During construction of the simulation setups, it was carefully considered that no leaflet introduced higher lipid densities that could result in artificial displacement effects. The systems comprised 0.8-1.65 million atoms and were modeled using the CHARMM36 force field (82) in combination with in-house DFT-based parameters (83) of the co-factors. Protonation states were established based on electrostatic calculations with Monte Carlo sampling techniques using APBS/Karlsberg+ (84-86).

Revised version: Atomistic molecular dynamics (aMD) simulations of the ovine Complex I (CI), Complex III2 (CIII2), and the I-III2 supercomplex (SCI/III2) were performed in a POPC:POPE:cardiolipin (2:2:1) membrane containing 5 mol% QH2 / Q (1:1 ratio). Cardiolipin was modeled as tetraoleoyl cardiolipin (18:1,18:1,18:1,18:1) with a headgroup modeled in a singly protonated state (with Qtot=-1). The system was solvated with TIP3P water molecules and 150 mM NaCl. The simulations were performed in different ligand states (apo, Q/QH2 bound, Table S1) at T=310 K and p=1 bar using NAMD2.14 (89) with a 2 fs integration timestep and long-range electrostatic interactions treated using the Particle Mesh Ewald approach. During construction of the simulation setups, it was carefully considered that no leaflet introduced higher lipid densities that could result in artificial displacement effects. The systems comprised 0.8-1.65 million atoms and were modeled using the CHARMM36 force field (90) in combination with in-house DFT-based parameters (91) of the co-factors. Protonation states were established based on electrostatic calculations with Monte Carlo sampling techniques using APBS/Karlsberg+ (92-94).

Earlier version: Molecular dynamics (MD) simulations of the ovine CI were conducted based on a cryo-EM structure (PDB ID: 6ZKC (52)), with simulations performed for both the quinol-bound and apo forms. Unresolved regions in the supernumerary subunits NDUFA7 and NDUFB6 were modeled using ColabFold (87). The system was equilibrated for 100 ns in the QH2-bound state, which was then used to propagate both states for 2×0.5 µs each.

Revised version: Molecular dynamics (MD) simulations of the ovine CI were conducted based on a cryo-EM structure (PDB ID: 6ZKC (54)), with simulations performed for both the quinol-bound and apo forms. Unresolved regions in the supernumerary subunits NDUFA7 and NDUFB6 were modeled using ColabFold (95). The system was equilibrated for 100 ns in the QH2-bound state, which was then used to propagate both states for 2×0.5 µs each.

The ovine Complex III2 was modeled based on the cryo-EM structure (PDB ID: 6Q9E (40)), with simulations also performed in different ligand states (Table S1). After equilibration of the system for 100 ns with the distal and proximal Q sites occupied, the system was propagated in each state for 2×0.5 µs.

A fully atomistic model of the SC was constructed based on the high-resolution structures of CI and CIII2, that were merged into the experimental structure of the supercomplex (PDB ID: 6QBX (40)). After equilibration of the system for 100 ns with all Q sites occupied in CI and CIII2, the complete system was propagated for 2×0.5 µs in each state (see Table S1).

Membrane systems with 1:1 POPC/POPE or CDL were constructed using CharmmGUI with a membrane area of 80×80 Å2, a hydration layer of 45 Å, and NaCl or KCl concentrations of 150 mM. See Table S2 for a detailed description of the membrane compositions. See the SI Appendix for a detailed description of the simulations and analysis.

Coarse-grained MD simulations

Earlier version: Coarse-grained MD (cgMD) simulation models of the ovine CI, CIII2, and the SCI1/III2 were created based on the atomistic models using the MARTINI3 force field (88) and Gromacs (89). All simulations were constructed with identical amounts of lipid molecules, a 2:2:1 POPC:POPE:cardiolipin ratio, and a 3 mol% mixture of quinone/quinol embedded in a simulation system with dimensions 47×47×31 nm3 (for the SC, see Table S1). All simulations were carried out at 310 K in an NPT ensemble using the velocity rescaling thermostat (90, 91), the Parrinello-Rahman barostat (92) and a 20 fs timestep using GROMACS (89). The protein structures were stabilized with elastic networks on backbone beads using a cutoff distance of 0.9 nm with a force constant of 500 kJ mol−1 nm−2. The elastic network was also applied between residues of different subunits, but not between CI and CIII2, whilst cgMD parameters for all cofactors were also developed. Two replicas were run for CI and CIII2 (2×50 µs each), and for the SCI/III2 (2×50 µs), as well as 23 µs cgMD simulations of the membrane with the same number of lipid and Q/QH2 molecules as in the protein simulations. Additional 5 µs cgMD simulations of the membrane systems, as well as for all protein models, were performed with longer lipids (0.44 nm, 0.50 nm, and 0.53 nm instead of 0.47 nm).

Revised version: Coarse-grained MD (cgMD) simulation models of the ovine CI, CIII2, and the SCI1/III2 were created based on the atomistic models using the MARTINI3 force field (96) and Gromacs (97). All simulations were constructed with identical amounts of lipid molecules, a 2:2:1 POPC:POPE:cardiolipin ratio, and a 3 mol% mixture of quinone/quinol embedded in a simulation system with dimensions 47×47×31 nm3 (for the SC, see Table S1). All simulations were carried out at 310 K in an NPT ensemble using the velocity rescaling thermostat (98, 99), the Parrinello-Rahman barostat (100) and a 20 fs timestep using GROMACS (97). The protein structures were stabilized with elastic networks on backbone beads using a cutoff distance of 0.9 nm with a force constant of 500 kJ mol−1 nm−2. The elastic network was also applied between residues of different subunits, but not between CI and CIII2, whilst cgMD parameters for all cofactors were also developed. Two replicas were run for CI and CIII2 (2×50 µs each), and for the SCI/III2 (2×50 µs), as well as 23 µs cgMD simulations of the membrane with the same number of lipid and Q/QH2 molecules as in the protein simulations. Additional 5 µs cgMD simulations of the membrane systems, as well as for all protein models, were performed with longer lipids (0.44 nm, 0.50 nm, and 0.53 nm instead of 0.47 nm).

Lattice model of SC formation

A lattice model of the CI and CIII2 was constructed (Fig. 4A,B) by modeling the OXPHOS proteins in unique grid positions on a 2D N×N lattice, with N=[6,15]. Depending on the relative orientation, the protein-protein interaction was described by specific interactions (giving rise to the energetic contribution Especific < 0) and non-specific interactions (Enon-specific > 0), whereas the membrane-protein interaction determined the strain energy of the membrane (Estrain), based on the number of neighboring “lipid” occupied grids that are in contact with proteins (Fig. 4A). The interaction between the lipids was indirectly accounted for by the background energy of the model. The proteins could occupy four unique orientations on a grid (gi=[North, East, South, West]). The states and their respective energies that the system can visit are summarized in Table S6.

Added in the revised version: The total Hamiltonian of the lattice model can be written as,

Added in the revised version: With

Earlier version: The total Hamiltonian of the lattice model can be written as, with The gigj -term assigns a specific energy contribution when the OXPHOS complexes are in adjacent lattice points only in a correct orientation (modeling a specific non-covalent interaction between the complexes such as the Arg29FB4-Asp260C1/Glu259C1 interaction between CI and CIII2). The didj -term assigns a non-specific interaction for the OXPHOS complexes when they are in adjacent lattice points, but in a “wrong” orientation relative to each other to form a specific interaction. The δij term introduces a strain into all lattice points surrounding an OXPHOS complex, mimicking the local membrane perturbation effects observed in our molecular simulations.

Revised version: The gigj -term assigns a specific energy contribution when the OXPHOS complexes are in adjacent lattice points only in a correct orientation (modeling a specific non-covalent interaction between the complexes such as the Arg29FB4-Asp260C1/Glu259C1 interaction between CI and CIII2). The didj -term assigns a non-specific interaction for the OXPHOS complexes when they are in adjacent lattice points, but in a “wrong” orientation relative to each other to form a specific interaction. The δij term introduces a strain into all lattice points surrounding an OXPHOS complex, mimicking the local membrane perturbation effects observed in our molecular simulations.

Earlier version: This leads to the partition function, where wi is the degeneracy of the state, modeling that the SC and OXPHOS proteins can reside at any lattice position of the membrane, and where β=1/kBT (kB, Boltzmann’s constant; T, temperature). The probability of a given state i was calculated as, with the free energy (G) defined as, The conformational landscape was sampled by Monte Carlo (MC) using 107 MC iterations with 100 replicas (see SI Appendix). Temperature effects were modeled by varying β, and the effect of different protein-to-lipid ratios by increasing the grid area. The simulation details can be found in Table S7.

Revised version: This leads to the partition function, where wi is the degeneracy of the state, modeling that the SC and OXPHOS proteins can reside at any lattice position of the membrane, and where β=1/kBT (kB, Boltzmann’s constant; T, temperature). The probability of a given state i was calculated as, with the free energy (G) defined as,

Added in the revised version: The conformational landscape was sampled by Monte Carlo (MC) using 107 MC iterations with 100 replicas (see SI Appendix). Temperature effects were modeled by varying β, and the effect of different protein-to-lipid ratios by increasing the grid area. The simulation details can be found in Table S7.

Statistical model of the membrane distribution in the IMM

Based on the experimental protein copy numbers (cf. Ref (7)) and average protein areas, the proteins in the IMM were modeled as randomly positioned circles on a membrane square with a side length of 163 nm (see Table S5). The system was minimized according to the Hamiltonian, where dij is the distance between two proteins i and j with respective radii, ri and rj, while the force constant k, were introduced to avoid steric clashes and set to 1 kcal mol-1 nm-1. This Hamiltonian corresponds to a two-dimensional system consisting of rigid, non-deformable circular particles.